---
title: 'video pideo fideo'
description: 'just testing again if the video exercise works, what can I do? I wont go over the audio until this is solved. What am I missing? I am not doing anything funny.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise, the whole point is to make the video work.

`@instructions`
This is an example exercise, the whole point is to make the video work.

`@hint`
solo te pido que no te r..... y con eso te valio

`@pre_exercise_code`
```{r}
#say hi by clicking submit answer
print("hi")
```

`@sample_code`
```{r}
#say hi by clicking submit answer
print("hi")
```

`@solution`
```{r}
print("hi")
```

`@sct`
```{r}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: e58495712e
xp: 50
```

`@projector_key`
623aee2ace60ad1f90cffdc51ac58e63
